<a href="https://colab.research.google.com/github/jacob-torres/DS-Unit-2-Kaggle-Challenge/blob/master/module1-decision-trees/Jacob_Torres_DS_module_project_221.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Lambda School Data Science

*Unit 2, Sprint 2, Module 1*

---

```python
import sys

# If you're on Colab:
if 'google.colab' in sys.modules:
    !pip install category_encoders==2.*
    !pip install pandas-profiling==2.*

DATA_PATH = 'https://raw.githubusercontent.com/LambdaSchool/DS-Unit-2-Kaggle-Challenge/master/data/'
```

Output:
```
Requirement already satisfied: category_encoders==2.* in /usr/local/lib/python3.6/dist-packages (2.2.2)
Requirement already satisfied: numpy>=1.14.0 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (1.18.5)
Requirement already satisfied: scipy>=1.0.0 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (1.4.1)
Requirement already satisfied: statsmodels>=0.9.0 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (0.10.2)
Requirement already satisfied: pandas>=0.21.1 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (1.1.4)
Requirement already satisfied: patsy>=0.5.1 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (0.5.1)
Requirement already satisfied: scikit-learn>=0.20.0 in /usr/local/lib/python3.6/dist-packages (from category_encoders==2.*) (0.22.2.post1)
Requirement already satisfied: pytz>=2017.2 in /usr/local/lib/python3.6/dist-packages (from pandas>=0.21.1->category_encoders=
```

# Decision Trees

## Kaggle

**Task 1:** [Sign up for a Kaggle account](https://www.kaggle.com/), if you don’t already have one. Go to our Kaggle InClass competition website (the URL is in Slack). Go to the Rules page. Accept the rules of the competition. Notice that the Rules page also has instructions for the Submission process. The Data page has feature definitions.

## Wrangle Data

**Task 2:** Add to the code below so that `id` is set as the index for `df`.

<!-- grade_id: cell-e2eb2544508f89ef, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = pd.merge(pd.read_csv(DATA_PATH + 'waterpumps/train_features.csv'),
              pd.read_csv(DATA_PATH + 'waterpumps/train_labels.csv'))
df = df.set_index('id')

df.head()
```

Output:
```
amount_tsh date_recorded  ... waterpoint_type_group    status_group
id                               ...                                      
69572      6000.0    2011-03-14  ...    communal standpipe      functional
8776          0.0    2013-03-06  ...    communal standpipe      functional
34310        25.0    2013-02-25  ...    communal standpipe      functional
67743         0.0    2013-01-28  ...    communal standpipe  non functional
19728         0.0    2011-07-13  ...    communal standpipe      functional

[5 rows x 40 columns]
```

**Test 2**

<!-- grade_id: cell-caa83cb6363d0cd7, points: 1, grade: true, solution: false -->
```python
'''Task 2 Testing'''

assert isinstance(df, pd.DataFrame), 'Have you created the DataFrame `df`?'
assert df.shape == (59400, 40), '`df` is the wrong shape. Did you set the index to `id`?'
assert 69572 in df.index, 'The index for `df` is has the wrong values. Did you set the index to `id`?'
```

### Exploratory Data Analysis:

## Split Data

**Task 3:** Create your target vector `y` and feature matrix `X`.

<!-- grade_id: cell-c94bceb606d02353, points: 0, grade: false, solution: true -->
```python
target = 'status_group'
X = df.drop(columns=target)
y = df[target]
```

**Test 3**

<!-- grade_id: cell-d0b6777daf75d805, points: 1, grade: true, solution: false -->
```python
'''Task 3 Testing'''

assert isinstance(X, pd.DataFrame), '`X` is the wrong data type.'
assert isinstance(y, pd.Series), '`y` is the wrong data type.'
assert y.shape == (59400,), '`y` is the wrong shape.'
assert X.shape == (59400,39), '`X` is the wrong shape.'
```

We already have a test set for this model, the `test.csv` that you'll use to make the predictions you'll upload to Kaggle. However, since our competition only allows for 2 submissions per day, we need a way to estimate our training error so that we can quickly iterate and improve our model. We can achieve this by creating a validation set from the data we have.

**Task 4:** Split `X` and `y` into training and validation sets. Your validation set should be 20% of the data you have. You should have four variables: `X_train`, `X_val`, `y_train`, and `y_val`. 

<!-- grade_id: cell-018b7c0deecc89dc, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

X_train, X_val, y_train, y_val = train_test_split(
    X, y, test_size=0.2, random_state=122995
)
```

**Test 4**

<!-- grade_id: cell-b1436df7d5901b26, points: 1, grade: true, solution: false -->
```python
'''Task 4 Testing'''

assert X_train.shape == (47520, 39), '`X_train` is the wrong shape.'
assert X_val.shape == (11880, 39), '`X_val` is the wrong shape.'
assert y_train.shape == (47520,), '`y_train` is the wrong shape.'
assert y_val.shape == (11880,), '`y_val` is the wrong shape.'
```

## Establish Baseline

**Task 5:** This is a **classification** problem, so you need to establish the baseline accuracy for your training set. Find the majority class for `y_train` and calculate the percentage of labels in `y_train` belonging to that class. Assign your answer to the variable name `baseline_acc`.

<!-- grade_id: cell-c6c63e2b6ff9e101, points: 0, grade: false, solution: true -->
```python
baseline_acc = 0.0
counts = y_train.value_counts(normalize=True)
for val in counts:
  if val >= baseline_acc:
    baseline_acc = val

print('Baseline Accuracy:', baseline_acc)
```

Output:
```
Baseline Accuracy: 0.5425084175084175
```

**Task 5**

<!-- grade_id: cell-a0c23a1103429de2, points: 1, grade: true, solution: false -->
```python
'''Task 5 Testing'''

assert isinstance(baseline_acc, float), '`baseline_acc` should be a `float`.'
assert 0.0 <= baseline_acc <= 1.0, '`baseline_acc` is a score that should be between 0 and 1.'
```

## Build Model

**Task 6:** Create a model named `model` and train it with your training data. Your model should be a pipeline with (a) transformers that you think are appropriate to this dataset and (b) a `DecisionTreeClassifier` as your predictor. **Tip:** How can you transform categorical features and missing values in order to train your model?

<!-- grade_id: cell-0c7c4d20b0f0fc70, points: 0, grade: false, solution: true -->
```python
from category_encoders import OrdinalEncoder
from sklearn.impute import SimpleImputer
from sklearn.pipeline import make_pipeline
from sklearn.tree import DecisionTreeClassifier

model = make_pipeline(OrdinalEncoder(), SimpleImputer(), DecisionTreeClassifier())
model.fit(X_train, y_train);
```

Output:
```
/usr/local/lib/python3.6/dist-packages/statsmodels/tools/_testing.py:19: FutureWarning: pandas.util.testing is deprecated. Use the functions in the public API at pandas.testing instead.
  import pandas.util.testing as tm
```

**Test 6**

<!-- grade_id: cell-4edc1a13268269ac, points: 1, grade: true, solution: false -->
```python
'''Task 6 Testing'''

assert len(model) > 1, 'Your model pipeline should have multiple steps.'
assert isinstance(model[-1], DecisionTreeClassifier), 'Your pipline should end in a `DecisionTreeClassifier`.'
assert hasattr(model, 'classes_'), 'Have you fit your model?'
```

## Check Metrics

**Task 7:** Calculate the training and validation accuracy of your model, and assign them to the valiables `training_acc` and `validation_acc`, respectively. Your validation accuracy should be greater than your baseline accuracy.

<!-- grade_id: cell-87a17042f6131ba5, points: 0, grade: false, solution: true -->
```python
y_pred_train = model.predict(X_train)
y_pred_val = model.predict(X_val)
training_acc = model.score(X_train, y_train)
validation_acc = model.score(X_val, y_val)
print('Training Accuracy:', training_acc)
print('Validation Accuracy:', validation_acc)
```

Output:
```
Training Accuracy: 0.9999789562289563
Validation Accuracy: 0.7126262626262626
```

**Test 7**

<!-- grade_id: cell-5b2575d28995d5be, points: 1, grade: true, solution: false -->
```python
'''Task 7 Testing'''

assert isinstance(training_acc, float)
assert isinstance(validation_acc, float)
assert 0.0 <= training_acc <= 1.0
assert 0.0 <= validation_acc <= 1.0
```

## Create Kaggle Submission

**Task 8:** Load `'waterpumps/test.csv'` into a DataFrame named `X_test`. Generate a list of predictions, and then put them into a DataFrame `submission`. Be sure that `submission` has the same index as `X_test` and that the column name for your predictions is `'status_group'`.  

<!-- grade_id: cell-f449c26c27917323, points: 0, grade: false, solution: true -->
```python
X_test = pd.read_csv(DATA_PATH + 'waterpumps/test_features.csv', index_col='id')
y_test = model.predict(X_test)

submission = pd.DataFrame(
    data=y_test, index=X_test.index, columns=['status_group']
)

submission.head()
```

Output:
```
status_group
id                   
50785      functional
51630  non functional
17168      functional
45559  non functional
49871      functional
```

**Test 8**

<!-- grade_id: cell-7184784b1c67a971, points: 1, grade: true, solution: false -->
```python
'''Task 8 Testing'''

assert isinstance(submission, pd.DataFrame), '`submission` should be a DataFrame.'
assert len(submission) == 14358, '`submission` should have 14358 rows.'
```

**Task 9 (`stretch goal`):** Save `submission` as a csv file using [`.to_csv()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.to_csv.html) and submit it to our Kaggle competition.

<!-- grade_id: cell-6a87d8169b9c48ea, points: 0, grade: false, solution: true -->
```python
from google.colab import files
submission_csv = submission.to_csv('waterpumps_submission_1.csv')
files.download('waterpumps_submission_1.csv')
```

Output:
```
<IPython.core.display.Javascript object>
<IPython.core.display.Javascript object>
```

## Explain

**Task 10 (`stretch goal`):** Make a horizontal barchart of the 10 most important features for your model.

<!-- grade_id: cell-297ec1119fbbaad5, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt

# YOUR CODE HERE
raise NotImplementedError()
```